# Aula 12 - Prática Guiada em Sala
**Curso:** Programação para Ciência de Dados  
**Tópico:** Dados Multidimensionais (Tensores 3D+)  
**Data:** 13 de Novembro 2025  

---

## Objetivos da Prática

Nesta sessão prática, você irá:

1. Manipular tensores 3D e 4D (arrays multidimensionais)
2. Dominar o parâmetro `axis` em operações de agregação
3. Aplicar indexação avançada (fancy e boolean)
4. Utilizar broadcasting eficientemente
5. Normalizar dados multidimensionais
6. Trabalhar com diferentes tipos de dados 3D+
7. Construir pipeline completo de processamento
8. Otimizar código para performance

---

## Dataset: MNIST Digits

Utilizaremos o dataset **MNIST** (Modified National Institute of Standards and Technology):
- 1797 imagens de dígitos manuscritos (0-9)
- Formato: 8×8 pixels em grayscale
- Shape: (1797, 8, 8) - tensor 3D
- Valores: 0-16 (intensidade de pixels)


```python
# === CONFIGURAÇÃO INICIAL ===

!pip install --upgrade pip --quiet
!pip cache purge
!pip install --upgrade otter-grader --quiet
!mkdir -p tests

print("Configuração concluída!")
```

```python
%%writefile tests/p1.py
OK_FORMAT = True

test = {
    "name": "p1",
    "points": 1,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
>>> # Teste 1: Retorna dicionário
>>> import numpy as np
>>> arr = np.random.rand(100, 8, 8)
>>> resultado = analise_shape(arr)
>>> isinstance(resultado, dict)
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 2: Chaves corretas
>>> import numpy as np
>>> arr = np.random.rand(10, 5, 3)
>>> resultado = analise_shape(arr)
>>> set(resultado.keys()) == {'ndim', 'shape', 'size', 'dtype'}
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 3: Valores corretos
>>> import numpy as np
>>> arr = np.ones((5, 4, 3))
>>> resultado = analise_shape(arr)
>>> resultado['ndim'] == 3 and resultado['size'] == 60
True
""",
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

```python
%%writefile tests/p2.py
OK_FORMAT = True

test = {
    "name": "p2",
    "points": 1.5,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
>>> # Teste 1: Shape correto axis=0
>>> import numpy as np
>>> arr = np.random.rand(10, 5, 3)
>>> resultado = agregar_axis(arr, axis=0, operacao='mean')
>>> resultado.shape == (5, 3)
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 2: Operação correta
>>> import numpy as np
>>> arr = np.ones((5, 4, 3))
>>> resultado = agregar_axis(arr, axis=1, operacao='sum')
>>> np.allclose(resultado, np.ones((5, 3)) * 4)
True
""",
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

```python
%%writefile tests/p3.py
OK_FORMAT = True

test = {
    "name": "p3",
    "points": 1.5,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
>>> # Teste 1: Shape preservado
>>> import numpy as np
>>> arr = np.random.rand(50, 10, 10)
>>> resultado = selecionar_indices(arr, [0, 5, 10], [2, 3, 4], [1, 2, 3])
>>> resultado.shape == (3,)
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 2: Boolean indexing correto
>>> import numpy as np
>>> arr = np.random.rand(100, 8, 8)
>>> resultado = filtrar_por_media(arr, threshold=0.5)
>>> resultado.ndim == 3
True
""",
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

```python
%%writefile tests/p4.py
OK_FORMAT = True

test = {
    "name": "p4",
    "points": 1.5,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
>>> # Teste 1: Normalização global
>>> import numpy as np
>>> arr = np.array([[[0, 10], [5, 15]]])
>>> resultado = normalizar_tensor(arr, metodo='global')
>>> bool(np.isclose(resultado.min(), 0) and np.isclose(resultado.max(), 1))
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 2: Shape preservado
>>> import numpy as np
>>> arr = np.random.rand(20, 8, 8)
>>> resultado = normalizar_tensor(arr, metodo='por_amostra')
>>> resultado.shape == arr.shape
True
""",
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

```python
%%writefile tests/p5.py
OK_FORMAT = True

test = {
    "name": "p5",
    "points": 1,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
>>> # Teste 1: Broadcasting funciona
>>> import numpy as np
>>> arr = np.random.rand(10, 5, 3)
>>> resultado = aplicar_broadcasting(arr, operacao='subtrair_media')
>>> resultado.shape == arr.shape
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 2: Média próxima de zero
>>> import numpy as np
>>> arr = np.random.rand(50, 10, 10)
>>> resultado = aplicar_broadcasting(arr, operacao='subtrair_media')
>>> bool(np.abs(resultado.mean()) < 1e-10)
True
""",
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

```python
%%writefile tests/p6.py
OK_FORMAT = True

test = {
    "name": "p6",
    "points": 1,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
>>> # Teste 1: Retorna dicionário
>>> import numpy as np
>>> arr = np.random.rand(100, 5, 3)
>>> resultado = analisar_series_temporal(arr)
>>> isinstance(resultado, dict)
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 2: Shapes corretos
>>> import numpy as np
>>> arr = np.random.rand(50, 3, 2)
>>> resultado = analisar_series_temporal(arr)
>>> resultado['media_temporal'].shape == (3, 2)
True
""",
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

```python
%%writefile tests/p7.py
OK_FORMAT = True

test = {
    "name": "p7",
    "points": 1.5,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
>>> # Teste 1: Retorna dicionário
>>> import numpy as np
>>> from sklearn.datasets import load_digits
>>> digits = load_digits()
>>> X, y = digits.images, digits.target
>>> resultado = pipeline_mnist(X, y, digitos=[0, 1, 2])
>>> isinstance(resultado, dict)
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 2: Imagens médias shape correto
>>> import numpy as np
>>> from sklearn.datasets import load_digits
>>> digits = load_digits()
>>> X, y = digits.images, digits.target
>>> resultado = pipeline_mnist(X, y, digitos=[0, 1])
>>> resultado['imagens_medias'].shape == (2, 8, 8)
True
""",
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

```python
%%writefile tests/p8.py
OK_FORMAT = True

test = {
    "name": "p8",
    "points": 1,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
>>> # Teste 1: Versão vetorizada mais rápida
>>> import numpy as np
>>> arr = np.random.rand(1000, 10, 10)
>>> tempo_loop, tempo_vec = comparar_performance(arr)
>>> tempo_vec < tempo_loop
True
""",
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
>>> # Teste 2: Resultados idênticos
>>> import numpy as np
>>> arr = np.random.rand(50, 5, 5)
>>> resultado_loop = processar_com_loop(arr)
>>> resultado_vec = processar_vetorizado(arr)
>>> np.allclose(resultado_loop, resultado_vec)
True
""",
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

```python
# === CARREGAR OTTER GRADER ===

import otter
grader = otter.Notebook()

print("Otter Grader carregado com sucesso!")

```

```python
# === IMPORTAR BIBLIOTECAS ===

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.datasets import load_digits
import time

# Configurações de visualização
plt.style.use('default')
plt.rcParams['figure.figsize'] = (12, 6)

# Seed para reprodutibilidade
np.random.seed(42)

print("Bibliotecas importadas!")
print(f"NumPy version: {np.__version__}")
print(f"Pandas version: {pd.__version__}")
```

```python
# === CARREGAR DATASET MNIST ===

digits = load_digits()
X = digits.images  # (1797, 8, 8)
y = digits.target  # (1797,)

print("=== DATASET MNIST CARREGADO ===")
print(f"Shape das imagens (X): {X.shape}")
print(f"Shape dos labels (y): {y.shape}")
print(f"Número de classes: {len(np.unique(y))}")
print(f"Range de valores: [{X.min()}, {X.max()}]")
print(f"Dtype: {X.dtype}")

# Visualizar algumas imagens
fig, axes = plt.subplots(2, 5, figsize=(12, 5))
for i in range(10):
    ax = axes[i // 5, i % 5]
    ax.imshow(X[i], cmap='gray')
    ax.set_title(f'Label: {y[i]}')
    ax.axis('off')
plt.suptitle('Primeiras 10 Imagens do MNIST', fontsize=14)
plt.tight_layout()
plt.show()

print("\nDataset pronto para uso!")
```

---

# PRÁTICA 1: Análise de Shape

## Objetivo
Entender e analisar as propriedades básicas de tensores multidimensionais.

## Tarefa
Implemente a função `analise_shape()` que recebe um tensor e retorna um dicionário com:
- `'ndim'`: número de dimensões
- `'shape'`: tupla com tamanho de cada dimensão
- `'size'`: total de elementos
- `'dtype'`: tipo de dados

## Dicas
```python
arr.ndim     # Número de dimensões
arr.shape    # Shape (tupla)
arr.size     # Total de elementos
arr.dtype    # Tipo de dados
```

## Retorno Esperado
Dicionário com 4 chaves.


```python
def analise_shape(arr):
    """
    Analisa propriedades básicas de um tensor.

    Args:
        arr (np.ndarray): Tensor de qualquer dimensionalidade

    Returns:
        dict: Dicionário com propriedades do tensor
    """
    # === SEU CÓDIGO AQUI ===

    info = {
        'ndim': arr.ndim,
        'shape': arr.shape,
        'size': arr.size,
        'dtype': arr.dtype
    }

    return info

    # === FIM DO SEU CÓDIGO ===
```

```python
# Testar função
print("=== ANÁLISE DO MNIST ===")
info_mnist = analise_shape(X)

for chave, valor in info_mnist.items():
    print(f"{chave}: {valor}")

# Interpretar dimensões
print("\n=== INTERPRETAÇÃO ===")
print(f"Temos {info_mnist['shape'][0]} imagens")
print(f"Cada imagem tem {info_mnist['shape'][1]}x{info_mnist['shape'][2]} pixels")
print(f"Total de pixels no dataset: {info_mnist['size']}")

# Verificar teste
grader.check("p1")
```

---

# PRÁTICA 2: Operações com Axis

## Objetivo
Dominar o uso do parâmetro `axis` em operações de agregação.

## Tarefa
Implemente a função `agregar_axis()` que:
1. Recebe um tensor 3D, um `axis`, e uma `operacao` ('mean', 'sum', 'max', 'min')
2. Aplica a operação no eixo especificado
3. Retorna o resultado agregado

## Conceito Importante
- `axis=0`: agrega ao longo da primeira dimensão (dimensão desaparece)
- `axis=1`: agrega ao longo da segunda dimensão
- `axis=2`: agrega ao longo da terceira dimensão

## Exemplo
```python
arr = np.ones((10, 5, 3))
arr.mean(axis=0)  # (5, 3) - média ao longo das 10 amostras
arr.sum(axis=1)   # (10, 3) - soma ao longo das 5 features
```


```python
def agregar_axis(arr, axis, operacao='mean'):
    """
    Agrega tensor ao longo de um eixo específico.

    Args:
        arr (np.ndarray): Tensor 3D
        axis (int): Eixo para agregar (0, 1, ou 2)
        operacao (str): Operação ('mean', 'sum', 'max', 'min')

    Returns:
        np.ndarray: Tensor agregado
    """
    # === SEU CÓDIGO AQUI ===

    if operacao == 'mean':
        resultado = arr.mean(axis=axis)
    elif operacao == 'sum':
        resultado = arr.sum(axis=axis)
    elif operacao == 'max':
        resultado = arr.max(axis=axis)
    elif operacao == 'min':
        resultado = arr.min(axis=axis)
    else:
        raise ValueError(f"Operação '{operacao}' não suportada")

    return resultado

    # === FIM DO SEU CÓDIGO ===
```

```python
# Testar com MNIST
print("=== AGREGAÇÕES NO MNIST ===")
print(f"Shape original: {X.shape}\n")

# Agregar ao longo de axis=0 (média de todas as imagens)
media_todas_imagens = agregar_axis(X, axis=0, operacao='mean')
print(f"Media axis=0: {media_todas_imagens.shape}")
print("  → Média de cada posição de pixel através de todas as imagens\n")

# Agregar ao longo de axis=1 (média por linha)
media_linhas = agregar_axis(X, axis=1, operacao='mean')
print(f"Media axis=1: {media_linhas.shape}")
print("  → Média de cada linha de cada imagem\n")

# Agregar ao longo de axis=2 (média por coluna)
media_colunas = agregar_axis(X, axis=2, operacao='mean')
print(f"Media axis=2: {media_colunas.shape}")
print("  → Média de cada coluna de cada imagem\n")

# Visualizar média de todas as imagens
plt.figure(figsize=(6, 6))
plt.imshow(media_todas_imagens, cmap='gray')
plt.title('Imagem Média de Todo o Dataset', fontsize=14)
plt.colorbar(label='Intensidade Média')
plt.axis('off')
plt.tight_layout()
plt.show()

# Verificar teste
grader.check("p2")
```

---

# PRÁTICA 3: Indexação Avançada

## Objetivo
Aplicar fancy indexing e boolean indexing em tensores 3D.

## Tarefa
Implemente duas funções:

### 3.1: `selecionar_indices()` - Fancy Indexing
- Recebe tensor 3D e três listas de índices (para cada dimensão)
- Retorna elementos específicos usando fancy indexing

### 3.2: `filtrar_por_media()` - Boolean Indexing
- Recebe tensor 3D e um threshold
- Retorna apenas amostras cuja média espacial (axis=(1,2)) é maior que threshold

## Conceitos
```python
# Fancy indexing
arr[[1, 3, 5]]  # Seleciona linhas 1, 3, 5

# Boolean indexing
mask = arr.mean(axis=1) > 10
arr[mask]  # Seleciona onde máscara é True
```


```python
def selecionar_indices(arr, indices_0, indices_1, indices_2):
    """
    Seleciona elementos específicos usando fancy indexing.

    Args:
        arr (np.ndarray): Tensor 3D
        indices_0 (list): Índices para dimensão 0
        indices_1 (list): Índices para dimensão 1
        indices_2 (list): Índices para dimensão 2

    Returns:
        np.ndarray: Elementos selecionados
    """
    # === SEU CÓDIGO AQUI ===

    resultado = arr[indices_0, indices_1, indices_2]

    return resultado

    # === FIM DO SEU CÓDIGO ===


def filtrar_por_media(arr, threshold):
    """
    Filtra amostras por média espacial usando boolean indexing.

    Args:
        arr (np.ndarray): Tensor 3D (amostras, altura, largura)
        threshold (float): Valor de corte

    Returns:
        np.ndarray: Amostras filtradas
    """
    # === SEU CÓDIGO AQUI ===

    # Calcular média espacial de cada amostra
    medias = arr.mean(axis=(1, 2))

    # Criar máscara booleana
    mask = medias > threshold

    # Aplicar máscara
    resultado = arr[mask]

    return resultado

    # === FIM DO SEU CÓDIGO ===
```

```python
# Testar fancy indexing
print("=== FANCY INDEXING ===")
indices_imgs = [0, 10, 100]
indices_rows = [3, 4, 5]
indices_cols = [2, 3, 4]

pixels_selecionados = selecionar_indices(X, indices_imgs, indices_rows, indices_cols)
print(f"Shape resultante: {pixels_selecionados.shape}")
print(f"Valores: {pixels_selecionados}")
print("  → Selecionou pixels específicos de imagens específicas\n")

# Testar boolean indexing
print("=== BOOLEAN INDEXING ===")
threshold = 6.5
imagens_claras = filtrar_por_media(X, threshold)

print(f"Total de imagens originais: {X.shape[0]}")
print(f"Threshold: {threshold}")
print(f"Imagens acima do threshold: {imagens_claras.shape[0]}")
print(f"Percentual filtrado: {100 * imagens_claras.shape[0] / X.shape[0]:.1f}%\n")

# Visualizar comparação
fig, axes = plt.subplots(2, 5, figsize=(12, 5))
axes[0, 0].set_ylabel('Originais', fontsize=12, rotation=0, labelpad=40)
axes[1, 0].set_ylabel('Filtradas\n(claras)', fontsize=12, rotation=0, labelpad=40)

# Determine quantas imagens podemos exibir (mínimo entre 5 e disponíveis)
num_to_display = min(5, imagens_claras.shape[0])

for i in range(5):
    # Primeira linha: sempre mostra originais
    axes[0, i].imshow(X[i], cmap='gray')
    axes[0, i].set_title(f'Média: {X[i].mean():.1f}')
    axes[0, i].axis('off')

    # Segunda linha: só mostra se houver imagens filtradas suficientes
    if i < num_to_display:
        axes[1, i].imshow(imagens_claras[i], cmap='gray')
        axes[1, i].set_title(f'Média: {imagens_claras[i].mean():.1f}')
    else:
        # Exibe placeholder para slots vazios
        axes[1, i].text(0.5, 0.5, 'N/A',
                       ha='center', va='center',
                       transform=axes[1, i].transAxes,
                       fontsize=14, color='red')
    axes[1, i].axis('off')

plt.suptitle(f'Comparação: Originais vs Filtradas (threshold={threshold})', fontsize=14)
plt.tight_layout()
plt.show()

# Verificar teste
grader.check("p3")
```

---

# PRÁTICA 4: Normalização de Tensores

## Objetivo
Implementar diferentes estratégias de normalização em tensores 3D.

## Tarefa
Implemente a função `normalizar_tensor()` com três métodos:

### 1. Normalização Global
- Normaliza todo o tensor para [0, 1]
- Fórmula: `(arr - min) / (max - min)`

### 2. Normalização Por Amostra
- Normaliza cada amostra independentemente
- Usa broadcasting com `keepdims=True`

### 3. Padronização Global
- Z-score: `(arr - mean) / std`

## Conceitos de Broadcasting
```python
# keepdims preserva dimensões para broadcasting
means = arr.mean(axis=(1,2), keepdims=True)  # (n, 1, 1)
normalized = arr / means  # Broadcasting funciona!
```


```python
def normalizar_tensor(arr, metodo='global'):
    """
    Normaliza tensor 3D usando diferentes estratégias.

    Args:
        arr (np.ndarray): Tensor 3D (amostras, altura, largura)
        metodo (str): 'global', 'por_amostra', ou 'padronizar'

    Returns:
        np.ndarray: Tensor normalizado
    """
    # === SEU CÓDIGO AQUI ===

    if metodo == 'global':
        # Normalização global: min-max para [0, 1]
        min_val = arr.min()
        max_val = arr.max()
        resultado = (arr - min_val) / (max_val - min_val)

    elif metodo == 'por_amostra':
        # Normalização por amostra (broadcasting com keepdims)
        min_vals = arr.min(axis=(1, 2), keepdims=True)  # (n, 1, 1)
        max_vals = arr.max(axis=(1, 2), keepdims=True)
        resultado = (arr - min_vals) / (max_vals - min_vals + 1e-10)

    elif metodo == 'padronizar':
        # Padronização global (z-score)
        mean_val = arr.mean()
        std_val = arr.std()
        resultado = (arr - mean_val) / std_val

    else:
        raise ValueError(f"Método '{metodo}' não suportado")

    return resultado

    # === FIM DO SEU CÓDIGO ===
```

```python
# Testar diferentes normalizações
X_subset = X[:20].copy()

print("=== COMPARAÇÃO DE NORMALIZAÇÕES ===")
print(f"\nOriginal:")
print(f"  Range: [{X_subset.min():.2f}, {X_subset.max():.2f}]")
print(f"  Mean: {X_subset.mean():.2f}, Std: {X_subset.std():.2f}")

# Normalização global
X_norm_global = normalizar_tensor(X_subset, metodo='global')
print(f"\nNormalização Global:")
print(f"  Range: [{X_norm_global.min():.2f}, {X_norm_global.max():.2f}]")
print(f"  Mean: {X_norm_global.mean():.2f}, Std: {X_norm_global.std():.2f}")

# Normalização por amostra
X_norm_amostra = normalizar_tensor(X_subset, metodo='por_amostra')
print(f"\nNormalização Por Amostra:")
print(f"  Range global: [{X_norm_amostra.min():.2f}, {X_norm_amostra.max():.2f}]")
print(f"  Mean: {X_norm_amostra.mean():.2f}, Std: {X_norm_amostra.std():.2f}")
print(f"  Cada amostra tem seu próprio min=0 e max=1")

# Padronização
X_padronizado = normalizar_tensor(X_subset, metodo='padronizar')
print(f"\nPadronização (z-score):")
print(f"  Range: [{X_padronizado.min():.2f}, {X_padronizado.max():.2f}]")
print(f"  Mean: {X_padronizado.mean():.2e}, Std: {X_padronizado.std():.2f}")

# Visualizar comparação
fig, axes = plt.subplots(2, 4, figsize=(14, 6))
metodos = [
    (X_subset, 'Original'),
    (X_norm_global, 'Global'),
    (X_norm_amostra, 'Por Amostra'),
    (X_padronizado, 'Padronizado')
]

for row in range(2):
    for col, (data, titulo) in enumerate(metodos):
        idx = row * 2  # Mostrar imagens 0 e 2
        axes[row, col].imshow(data[idx], cmap='gray')
        axes[row, col].set_title(f'{titulo}\n[{data[idx].min():.2f}, {data[idx].max():.2f}]')
        axes[row, col].axis('off')

plt.suptitle('Comparação de Métodos de Normalização', fontsize=14)
plt.tight_layout()
plt.show()

# Verificar teste
grader.check("p4")
```

---

# PRÁTICA 5: Broadcasting Avançado

## Objetivo
Aplicar broadcasting para operações eficientes em tensores.

## Tarefa
Implemente a função `aplicar_broadcasting()` que:
1. Recebe tensor 3D e uma operação
2. Aplica operação usando broadcasting (SEM loops!)

### Operações Suportadas:
- `'subtrair_media'`: Subtrai média de cada amostra
- `'dividir_std'`: Divide pelo desvio padrão de cada amostra
- `'centralizar_pixel'`: Subtrai média de cada posição de pixel

## Conceito Chave
```python
# Broadcasting com keepdims
means = arr.mean(axis=(1,2), keepdims=True)  # (n, 1, 1)
centered = arr - means  # Broadcasting automático!
```


```python
def aplicar_broadcasting(arr, operacao='subtrair_media'):
    """
    Aplica operações usando broadcasting (vetorizado).

    Args:
        arr (np.ndarray): Tensor 3D (amostras, altura, largura)
        operacao (str): Tipo de operação

    Returns:
        np.ndarray: Tensor processado
    """
    # === SEU CÓDIGO AQUI ===

    if operacao == 'subtrair_media':
        # Subtrair média de cada amostra
        means = arr.mean(axis=(1, 2), keepdims=True)  # (n, 1, 1)
        resultado = arr - means

    elif operacao == 'dividir_std':
        # Dividir por desvio padrão de cada amostra
        stds = arr.std(axis=(1, 2), keepdims=True)  # (n, 1, 1)
        resultado = arr / (stds + 1e-10)  # Evitar divisão por zero

    elif operacao == 'centralizar_pixel':
        # Subtrair média de cada posição de pixel
        means = arr.mean(axis=0, keepdims=True)  # (1, h, w)
        resultado = arr - means

    else:
        raise ValueError(f"Operação '{operacao}' não suportada")

    return resultado

    # === FIM DO SEU CÓDIGO ===
```

```python
# Testar broadcasting
X_test = X[:50].copy()

print("=== OPERAÇÕES COM BROADCASTING ===")
print(f"Shape original: {X_test.shape}\n")

# Subtrair média
X_centered = aplicar_broadcasting(X_test, 'subtrair_media')
print("Subtrair Média:")
print(f"  Mean global original: {X_test.mean():.4f}")
print(f"  Mean global após: {X_centered.mean():.10f}")
print(f"  ✓ Cada amostra agora tem média ~0\n")

# Dividir por std
X_scaled = aplicar_broadcasting(X_test, 'dividir_std')
print("Dividir por Std:")
print(f"  Std médio das amostras: {X_scaled.std(axis=(1,2)).mean():.4f}")
print(f"  ✓ Escalou valores pelo desvio padrão\n")

# Centralizar por pixel
X_pixel_centered = aplicar_broadcasting(X_test, 'centralizar_pixel')
print("Centralizar por Pixel:")
print(f"  Media de cada posição de pixel após: {X_pixel_centered.mean(axis=0).mean():.10f}")
print(f"  ✓ Cada posição de pixel tem média ~0 através das amostras\n")

# Visualizar efeito de centralizar
fig, axes = plt.subplots(2, 3, figsize=(12, 8))
operacoes_vis = [
    (X_test, 'Original'),
    (X_centered, 'Centrado (por amostra)'),
    (X_pixel_centered, 'Centrado (por pixel)')
]

for row in range(2):
    for col, (data, titulo) in enumerate(operacoes_vis):
        idx = row * 10
        axes[row, col].imshow(data[idx], cmap='gray')
        axes[row, col].set_title(f'{titulo}\nMédia: {data[idx].mean():.2f}')
        axes[row, col].axis('off')

plt.suptitle('Efeitos do Broadcasting em Normalização', fontsize=14)
plt.tight_layout()
plt.show()

# Verificar teste
grader.check("p5")
```

---

# PRÁTICA 6: Análise de Séries Temporais Multivariadas

## Objetivo
Aplicar conceitos de tensores em dados de séries temporais.

## Contexto
Séries temporais multivariadas são tensores 3D:
- Shape: (tempo, variáveis, locais)
- Exemplo: (100 dias, 5 sensores, 3 horários)

## Tarefa
Implemente `analisar_series_temporal()` que retorna dicionário com:
- `'media_temporal'`: média ao longo do tempo (axis=0)
- `'media_por_dia'`: média de cada dia (axis=(1,2))
- `'variabilidade'`: desvio padrão ao longo do tempo
- `'dias_quentes'`: quantidade de dias onde média > threshold

## Conceito
Mesmas operações do MNIST, mas interpretação diferente!


```python
def analisar_series_temporal(arr, threshold=22.0):
    """
    Analisa série temporal multivariada.

    Args:
        arr (np.ndarray): Tensor 3D (tempo, variáveis, locais)
        threshold (float): Limite para dias quentes

    Returns:
        dict: Estatísticas da série
    """
    # === SEU CÓDIGO AQUI ===

    # Média temporal (média de cada sensor/horário ao longo do tempo)
    media_temporal = arr.mean(axis=0)

    # Média por dia (média de todos sensores/horários em cada dia)
    media_por_dia = arr.mean(axis=(1, 2))

    # Variabilidade temporal (quanto cada sensor varia ao longo do tempo)
    variabilidade = arr.std(axis=0)

    # Dias quentes (dias onde média está acima do threshold)
    dias_quentes = (media_por_dia > threshold).sum()

    resultado = {
        'media_temporal': media_temporal,
        'media_por_dia': media_por_dia,
        'variabilidade': variabilidade,
        'dias_quentes': dias_quentes
    }

    return resultado

    # === FIM DO SEU CÓDIGO ===
```

```python
# Criar dados sintéticos de temperatura
np.random.seed(42)
n_dias, n_sensores, n_horarios = 100, 5, 3

# Temperatura base + variação + tendência temporal
temp_base = 20
temp_data = np.random.randn(n_dias, n_sensores, n_horarios) * 3 + temp_base

# Adicionar tendência de aquecimento
tendencia = np.linspace(0, 5, n_dias)
temp_data += tendencia[:, np.newaxis, np.newaxis]

print("=== DADOS DE TEMPERATURA ===")
print(f"Shape: {temp_data.shape}")
print(f"  {n_dias} dias × {n_sensores} sensores × {n_horarios} horários")
print(f"Range: [{temp_data.min():.1f}°C, {temp_data.max():.1f}°C]")
print(f"Média global: {temp_data.mean():.1f}°C\n")

# Analisar série temporal
stats_temp = analisar_series_temporal(temp_data, threshold=22.0)

print("=== ANÁLISE DA SÉRIE TEMPORAL ===")
print(f"\nMédia Temporal (por sensor/horário):")
print(f"  Shape: {stats_temp['media_temporal'].shape}")
print(f"  Sensor mais quente: {stats_temp['media_temporal'].max():.1f}°C")
print(f"  Sensor mais frio: {stats_temp['media_temporal'].min():.1f}°C")

print(f"\nMédia Por Dia:")
print(f"  Shape: {stats_temp['media_por_dia'].shape}")
print(f"  Dia mais quente: {stats_temp['media_por_dia'].max():.1f}°C")
print(f"  Dia mais frio: {stats_temp['media_por_dia'].min():.1f}°C")

print(f"\nVariabilidade:")
print(f"  Sensor mais variável: std={stats_temp['variabilidade'].max():.2f}°C")
print(f"  Sensor mais estável: std={stats_temp['variabilidade'].min():.2f}°C")

print(f"\nDias Quentes (>22°C):")
print(f"  Total: {stats_temp['dias_quentes']} de {n_dias} dias")
print(f"  Percentual: {100 * stats_temp['dias_quentes'] / n_dias:.1f}%")

# Visualizar série temporal
fig, axes = plt.subplots(2, 1, figsize=(12, 8))

# Gráfico 1: Temperatura por dia
axes[0].plot(stats_temp['media_por_dia'], linewidth=2)
axes[0].axhline(22, color='r', linestyle='--', label='Threshold (22°C)')
axes[0].set_xlabel('Dia')
axes[0].set_ylabel('Temperatura Média (°C)')
axes[0].set_title('Evolução da Temperatura Média por Dia')
axes[0].legend()
axes[0].grid(True, alpha=0.3)

# Gráfico 2: Heatmap média temporal
im = axes[1].imshow(stats_temp['media_temporal'], cmap='RdYlBu_r', aspect='auto')
axes[1].set_xlabel('Horário')
axes[1].set_ylabel('Sensor')
axes[1].set_title('Temperatura Média por Sensor e Horário')
axes[1].set_xticks([0, 1, 2])
axes[1].set_xticklabels(['Manhã', 'Tarde', 'Noite'])
plt.colorbar(im, ax=axes[1], label='Temperatura (°C)')

plt.tight_layout()
plt.show()

# Verificar teste
grader.check("p6")
```

---

# PRÁTICA 7: Pipeline Completo MNIST

## Objetivo
Construir pipeline end-to-end de processamento de dados multidimensionais.

## Tarefa
Implemente `pipeline_mnist()` que:
1. Filtra apenas dígitos especificados
2. Normaliza imagens para [0, 1]
3. Calcula imagem média por classe
4. Calcula matriz de correlação entre classes
5. Retorna dicionário com todos os resultados

## Pipeline
```
Entrada → Filtrar → Normalizar → Agregar → Analisar → Resultados
```

## Retorno Esperado
Dicionário com:
- `'n_imagens_por_classe'`: array com contagens
- `'imagens_medias'`: tensor (n_classes, 8, 8)
- `'correlacao'`: matriz (n_classes, n_classes)


```python
def pipeline_mnist(X, y, digitos=[0, 1, 2]):
    """
    Pipeline completo de análise do MNIST.

    Args:
        X (np.ndarray): Imagens (n, 8, 8)
        y (np.ndarray): Labels (n,)
        digitos (list): Dígitos a analisar

    Returns:
        dict: Resultados da análise
    """
    # === SEU CÓDIGO AQUI ===

    # 1. Filtrar dígitos
    mask = np.isin(y, digitos)
    X_filtrado = X[mask]
    y_filtrado = y[mask]

    # 2. Normalizar para [0, 1]
    X_normalizado = (X_filtrado - X_filtrado.min()) / (X_filtrado.max() - X_filtrado.min())

    # 3. Calcular imagem média por classe
    n_imagens_por_classe = []
    imagens_medias = []

    for digito in digitos:
        mask_classe = y_filtrado == digito
        X_classe = X_normalizado[mask_classe]

        n_imagens_por_classe.append(X_classe.shape[0])
        imagens_medias.append(X_classe.mean(axis=0))

    n_imagens_por_classe = np.array(n_imagens_por_classe)
    imagens_medias = np.array(imagens_medias)

    # 4. Calcular matriz de correlação entre classes
    # Flatten cada imagem média para calcular correlação
    imagens_flat = imagens_medias.reshape(len(digitos), -1)
    correlacao = np.corrcoef(imagens_flat)

    # 5. Montar resultado
    resultado = {
        'n_imagens_por_classe': n_imagens_por_classe,
        'imagens_medias': imagens_medias,
        'correlacao': correlacao
    }

    return resultado

    # === FIM DO SEU CÓDIGO ===
```

```python
# Executar pipeline
digitos_analisar = [0, 1, 2, 3, 4]
resultados = pipeline_mnist(X, y, digitos=digitos_analisar)

print("=== RESULTADOS DO PIPELINE ===")
print(f"\nDígitos analisados: {digitos_analisar}")
print(f"\nImagens por classe:")
for i, digito in enumerate(digitos_analisar):
    print(f"  Dígito {digito}: {resultados['n_imagens_por_classe'][i]} imagens")

print(f"\nImagens médias shape: {resultados['imagens_medias'].shape}")
print(f"Matriz de correlação shape: {resultados['correlacao'].shape}")

# Visualizar imagens médias
n_digitos = len(digitos_analisar)
fig, axes = plt.subplots(1, n_digitos, figsize=(12, 3))

for i, digito in enumerate(digitos_analisar):
    axes[i].imshow(resultados['imagens_medias'][i], cmap='gray')
    axes[i].set_title(f'Dígito {digito}\n({resultados["n_imagens_por_classe"][i]} imgs)')
    axes[i].axis('off')

plt.suptitle('Imagens Médias por Classe', fontsize=14)
plt.tight_layout()
plt.show()

# Visualizar matriz de correlação
plt.figure(figsize=(8, 6))
im = plt.imshow(resultados['correlacao'], cmap='coolwarm', vmin=-1, vmax=1)
plt.colorbar(im, label='Correlação')
plt.xticks(range(n_digitos), digitos_analisar)
plt.yticks(range(n_digitos), digitos_analisar)
plt.xlabel('Dígito')
plt.ylabel('Dígito')
plt.title('Matriz de Correlação entre Imagens Médias')

# Adicionar valores na matriz
for i in range(n_digitos):
    for j in range(n_digitos):
        text = plt.text(j, i, f'{resultados["correlacao"][i, j]:.2f}',
                       ha="center", va="center", color="black", fontsize=10)

plt.tight_layout()
plt.show()

# Insights
print("\n=== INSIGHTS ===")
correlacao = resultados['correlacao']
np.fill_diagonal(correlacao, 0)  # Ignorar diagonal
max_corr_idx = np.unravel_index(correlacao.argmax(), correlacao.shape)
print(f"Dígitos mais similares: {digitos_analisar[max_corr_idx[0]]} e {digitos_analisar[max_corr_idx[1]]}")
print(f"Correlação: {correlacao[max_corr_idx]:.3f}")

# Verificar teste
grader.check("p7")
```

---

# PRÁTICA 8: Performance - Vetorização vs Loops

## Objetivo
Demonstrar a importância da vetorização para performance.

## Tarefa
Implemente três funções que fazem a MESMA coisa:

### 1. `processar_com_loop()` - Com loop Python (LENTO)
- Centraliza cada imagem subtraindo sua média
- Usa loop for

### 2. `processar_vetorizado()` - Vetorizado (RÁPIDO)
- Mesma operação, mas com broadcasting
- SEM loops

### 3. `comparar_performance()` - Compara tempos
- Executa ambas versões
- Retorna tempos e speedup

## Conceito
Vetorização pode ser 10-100x mais rápida que loops!



```python
def processar_com_loop(arr):
    """
    Centraliza imagens usando loop (LENTO).

    Args:
        arr (np.ndarray): Tensor 3D (n, h, w)

    Returns:
        np.ndarray: Tensor centralizado
    """
    # === SEU CÓDIGO AQUI ===

    resultado = np.zeros_like(arr, dtype=float)

    for i in range(arr.shape[0]):
        resultado[i] = arr[i] - arr[i].mean()

    return resultado

    # === FIM DO SEU CÓDIGO ===


def processar_vetorizado(arr):
    """
    Centraliza imagens usando broadcasting (RÁPIDO).

    Args:
        arr (np.ndarray): Tensor 3D (n, h, w)

    Returns:
        np.ndarray: Tensor centralizado
    """
    # === SEU CÓDIGO AQUI ===

    means = arr.mean(axis=(1, 2), keepdims=True)
    resultado = arr - means

    return resultado

    # === FIM DO SEU CÓDIGO ===


def comparar_performance(arr):
    """
    Compara performance de loop vs vetorizado.

    Args:
        arr (np.ndarray): Tensor 3D

    Returns:
        tuple: (tempo_loop, tempo_vetorizado)
    """
    # === SEU CÓDIGO AQUI ===

    # Medir tempo com loop
    start = time.time()
    _ = processar_com_loop(arr)
    tempo_loop = time.time() - start

    # Medir tempo vetorizado
    start = time.time()
    _ = processar_vetorizado(arr)
    tempo_vec = time.time() - start

    return tempo_loop, tempo_vec

    # === FIM DO SEU CÓDIGO ===
```

```python
# Benchmark com diferentes tamanhos
tamanhos = [100, 500, 1000, 1797]
resultados_bench = []

print("=== BENCHMARK: LOOP VS VETORIZADO ===")
print(f"{'Tamanho':<10} {'Loop (ms)':<12} {'Vetorizado (ms)':<17} {'Speedup'}")
print("-" * 55)

for n in tamanhos:
    arr_test = X[:n]
    tempo_loop, tempo_vec = comparar_performance(arr_test)
    speedup = tempo_loop / tempo_vec

    print(f"{n:<10} {tempo_loop*1000:<12.2f} {tempo_vec*1000:<17.2f} {speedup:.1f}x")
    resultados_bench.append((n, tempo_loop*1000, tempo_vec*1000, speedup))

# Verificar que resultados são idênticos
resultado_loop = processar_com_loop(X[:100])
resultado_vec = processar_vetorizado(X[:100])

print("\n=== VERIFICAÇÃO DE CORRETUDE ===")
print(f"Resultados idênticos: {np.allclose(resultado_loop, resultado_vec)}")
print(f"Diferença máxima: {np.abs(resultado_loop - resultado_vec).max():.10f}")

# Visualizar speedup
fig, axes = plt.subplots(1, 2, figsize=(14, 5))

# Gráfico 1: Tempos absolutos
tamanhos_arr = [r[0] for r in resultados_bench]
tempos_loop = [r[1] for r in resultados_bench]
tempos_vec = [r[2] for r in resultados_bench]

axes[0].plot(tamanhos_arr, tempos_loop, 'o-', label='Loop', linewidth=2, markersize=8)
axes[0].plot(tamanhos_arr, tempos_vec, 's-', label='Vetorizado', linewidth=2, markersize=8)
axes[0].set_xlabel('Número de Imagens')
axes[0].set_ylabel('Tempo (ms)')
axes[0].set_title('Tempo de Execução')
axes[0].legend()
axes[0].grid(True, alpha=0.3)

# Gráfico 2: Speedup
speedups = [r[3] for r in resultados_bench]
axes[1].plot(tamanhos_arr, speedups, 'o-', color='green', linewidth=2, markersize=8)
axes[1].axhline(1, color='red', linestyle='--', label='Sem ganho')
axes[1].set_xlabel('Número de Imagens')
axes[1].set_ylabel('Speedup (x vezes)')
axes[1].set_title('Ganho de Performance (Vetorizado vs Loop)')
axes[1].legend()
axes[1].grid(True, alpha=0.3)

plt.tight_layout()
plt.show()

print("\n=== CONCLUSÃO ===")
print(f"✓ Vetorização é {speedups[-1]:.1f}x mais rápida no dataset completo!")
print("✓ Sempre prefira operações vetorizadas a loops Python")
print("✓ Broadcasting + NumPy = Performance!")

# Verificar teste
grader.check("p8")
```

---

# 🎓 FINALIZAÇÃO

## Verificar Todas as Práticas

```python
# Executar todos os testes
grader.check_all()
```

---

# RESUMO DA PRÁTICA

### Fundamentos de Tensores
- Análise de shape, ndim, size, dtype
- Interpretação de dimensões em diferentes contextos

### Operações com Axis
- Agregação ao longo de diferentes eixos
- `axis=0`, `axis=1`, `axis=2` - cada um tem significado diferente
- keepdims para preservar dimensões

### Indexação Avançada
- Fancy indexing com listas de índices
- Boolean indexing com máscaras
- Seleção eficiente de subsets

### Normalização e Padronização
- Normalização global vs por amostra
- Broadcasting com keepdims
- Trade-offs de cada método

### Broadcasting
- Operações vetorizadas sem loops
- Regras de compatibilidade de shapes
- newaxis para adicionar dimensões

### Aplicações Práticas
- Séries temporais multivariadas
- Pipeline completo de análise
- Matriz de correlação

### Performance
- Vetorização 10-100x mais rápida que loops
- NumPy otimizado para operações em batch
- Sempre prefira broadcasting a loops


Tests for p1 (points: 1)
```
>>> import numpy as np
>>> arr = np.random.rand(100, 8, 8)
>>> resultado = analise_shape(arr)
>>> isinstance(resultado, dict)
True
```


Tests for p2 (points: 1.5)


Tests for p3 (points: 1.5)


Tests for p4 (points: 1.5)


Tests for p5 (points: 1)


Tests for p6 (points: 1)


Tests for p7 (points: 1.5)


Tests for p8 (points: 1)
